# Problem 1: Steady-state diffusive hillslope profiles (*12.5 points*)

In this exercise, you will be plotting a diffusive hillslope profile and exploring the effect of various parameters on the hillslope geometry. We'll start by creating a function for calculating the hillslope profiles, then explore the model behavior using different input parameter values and plotting the results.

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable names mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.

**Your score on this problem will be based on following criteria**:

- Properly converting the hillslope diffusion equation to a Python function
- Calculating hillslope profiles for a range of input values
- Plotting your results
- Including comments that explain what most lines in the code do
- Uploading your notebook to your GitHub repository for this week's exercise

## Part 1: Defining a hillslope profile function (*2 points*)

Your first task for this problem is to define a function to calculate the geometry of an *interfluve* (the region between two adjacent river valleys) shaped by hillslope diffusion. As we saw briefly in the [lesson for this week](https://introqg.github.io/site/lessons/L3/solving-diffusion.html), we can calculate this geometry using the diffusion equation for a hillslope profile that has the rivers located at $x = ±L$ and the apex of the interfluve at $x = 0$. If we assume the landscape is being uplifted at a rate $U$ and that the rivers elevations are fixed at $h = 0$, we find that the elevation of the interfluve is

\begin{equation}
  \Large
  h(x) = \frac{U}{2 \kappa}\left(L^{2} - x^{2} \right)
\end{equation}

where $\kappa$ is the sediment/soil diffusivity.

- Create a function called `hillslope_diffusion` that calculates the elevation of the interfluve $h$ as a function of horizontal distance $x$.
    - Your function should take the half-interfluve width `L`, the list of values at which the elevation $h$ should be calculated `x`, the landscape uplift rate `U`, and the sediment/soil diffusivity `kappa` as parameters.
    - **Note**: You can use a `for` loop inside your function, but it is not strictly necessary

```python
# Run this cell first to import the libraries we need and configure the plotting

# Import NumPy, Pandas, and Matplotlib
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt

# Inline plotting in the Jupyter notebooks
%matplotlib inline
```

<!-- grade_id: problem_1_ag_p1_hillslope_function, points: 0, grade: false, solution: true -->
```python
def hillslope_diffusion(L, x, U, kappa):
    """Calculates the elevation of a hillslope experiencing diffusion."""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_1_ag_p1_hillslope_function_test1, points: 1, grade: true, solution: false -->
```python
# These visible tests with fake data should work
from nose.tools import assert_equal

# Fake data
L = 10
x_test = np.linspace(-L, L, 51)
h_test = hillslope_diffusion(L=L, x=x_test, U=0.5, kappa=0.05)

# Print results
print('Hillslope elevation at x = 0 should be 500.0. My code calculates {0:.1f}.'.format(h_test[25]))

# Check that the function values are correct at the boundaries
assert_equal(round(h_test[0], 4), 0.0000)
assert_equal(round(h_test[-1], 4), 0.0000)

# Check some random values
assert_equal(round(h_test[35], 4), 420.0000)
assert_equal(round(h_test[10], 4), 320.0000)
```

<!-- grade_id: problem_1_ag_p1_hillslope_function_test2, points: 1, grade: true, solution: false -->
```python
# These visible tests with fake data should work

# Fake data 2
L = 1000
x_test2 = np.linspace(-L, L, 51)
h_test2 = hillslope_diffusion(L=L, x=x_test2, U=5, kappa=0.2)

# Print results
print('Hillslope elevation at x = 0 should now be 12500000.0. My code calculates {0:.1f}.'.format(h_test2[25]))

# Check that the function values are correct at the boundaries
assert_equal(round(h_test2[0], 4), 0.0000)
assert_equal(round(h_test2[-1], 4), 0.0000)

# Check some random values
assert_equal(round(h_test2[28], 4), 12320000.0)
assert_equal(round(h_test2[1], 4), 980000.0000)
```

## Part 2: Calculating a hillslope geometry (*1 point*)

With the `hillslope_diffusion` function defined, we can now move on to calculating the geometry of a hillslope.

- Calculate the elevation of a hillslope `h` for an interfluve that has two channels located 100m apart and a landscape uplift rate $U$ of 0.5 mm/a. Assume the diffusion coefficient $\kappa = 50 \times 10^{-3}$m<sup>2</sup>/a and the increment for values of $x$ is 1 meter.
    - For the variables passed to your `hillslope_diffusion` function, use the name `x` for the range of locations at which the elevation should be calculated, `L` for the inferfluve half-width, `U` for the uplift rate, and `kappa` for the diffusivity.
    - **Note**: Be careful to ensure your units are compatible. Distance and time units should be adjusted accordingly to ensure they work together.

<!-- grade_id: problem_1_ag_p2_hillslope, points: 0, grade: false, solution: true -->
```python
# Define points where the hillslope elevation should be calculated
x = None

# Calculate hillslope elevation
h = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_1_ag_p2_hillslope_test1, points: 0.5, grade: true, solution: false -->
```python
# This test should work
print('x at index 20 should be -30.0. My code calculates {0:.1f}.'.format(x[20]))

```

<!-- grade_id: problem_1_ag_p2_hillslope_test2, points: 0.5, grade: true, solution: false -->
```python
# This test should work
print('h at index 20 should be 8.0. My code calculates {0:.1f}.'.format(h[20]))

```

## Part 3: Plotting your results (*3 points*)

Your next task is to create a pandas DataFrame to store your calculated hillslope profile, make a plot of the hillslope profile, and add a few values as text on the plot.

- Create a pandas DataFrame called `hillslope` with the distance along the hillslope profile `x` as the index and a single data column called `'Default'` that contains the hillslope profile elevations `h`
- Plot the value of `'Default'` as a function of the distance along your hillslope profile
    - Plot this using a dashed black line
    - It may be helpful to again increase the figure size to something like 12 inches by 8 inches
- Add axis labels and a plot title
- Calculate
    1. the maximum slope of the profile as variable `max_slope`
    2. its value in degrees as variable `max_slope_deg`
    3. the maximum relief of the profile as variable `max_relief`
    4. the characteristic timescale of diffusion (or time constant) as variable `tau`
- Display the 4 values above as text on the plot using the `ax.text()` function
    - You can find useful information about these values in the [notes on solving the diffusion equation](https://introqg.github.io/site/lessons/L3/solving-diffusion.html) in this week's lesson
    - You may want to use [Google](https://www.google.fi) to look up the characteristic timescale of diffusion
- Add a figure caption in the Markdown cell below the plot explaining what it shows as if it was in a scientific publication

<!-- grade_id: problem_1_ag_p3_plot1, points: 0, grade: false, solution: true -->
```python
# Make DataFrame
hillslope = None

# Calculate the hillslope properties
max_slope = None
max_slope_deg = None
max_relief = None
tau = None

# Create the plot
ax = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_1_mg_p3_plot1, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: problem_1_ag_p3_plot1_test1, points: 0.5, grade: true, solution: false -->
```python
# This test should display the first 5 rows of your DataFrame
print('The first 5 rows of the hillslope DataFrame are\n', hillslope.head())

```

<!-- grade_id: problem_1_ag_p3_plot1_test2, points: 0.5, grade: true, solution: false -->
```python
# This test should display the max hillslope angle in degrees
print('The maximum hillslope angle in degrees should be 26.6. My calculated value is {0:.1f}.'.format(max_slope_deg))

```

## Part 4: Exploring the model (*4.5 points*)

Your next task is to explore the effect of different parameters on the hillslope geometry.

- Make an additional **three** plots for the three variations of the hillslope model parameters below, with each plot created in a separate Python cell
    - In each case be sure you only vary a single parameter from the original values given in Part 2
    - Store the resulting hillslope profiles as new columns in your `hillslope` DataFrame using the names below:
        1. Double the uplift rate `U` and store the hillslope profile in column `'Double uplift'`
        2. Double the diffusivity `kappa` and store the hillslope profile in column `'Double diffusivity'`
        3. Reduce the diffusivity `kappa` by half and store the hillslope profile in column `'Half diffusivity'`
    - To make it easy to compare to the first plot, plot the `'Default'` hillslope profile from Part 3 as a black dashed line and the value you are asked to calculate in a color of your choosing (not black, please :D)
    - As above, be sure to include axis labels, a title, and the calculated hillslope properties (e.g., slope, maximum relief, etc.) as text on the plot
- Add a figure caption in the Markdown cell beneath **each** plot explaining what it shows as if it was in a scientific publication.

### Part 4, plot 1: Double the uplift rate

Please create the plot and do the related calculations below.

<!-- grade_id: problem_1_ag_p4_plot2, points: 0, grade: false, solution: true -->
```python
# Double the uplift rate
hillslope['Double uplift'] = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_1_mg_p4_plot2, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: problem_1_ag_p4_plot2_test1, points: 0.5, grade: true, solution: false -->
```python
# This test should display the max hillslope angle in degrees
print('The maximum hillslope angle in degrees should now be 45.0. My calculated value is {0:.1f}.'.format(max_slope_deg))

```

### Part 4, plot 2: Double the diffusivity

Please create the plot and do the related calculations below.

<!-- grade_id: problem_1_ag_p4_plot3, points: 0, grade: false, solution: true -->
```python
# Double the diffusivity
hillslope['Double diffusivity'] = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_1_mg_p4_plot3, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: problem_1_ag_p4_plot3_test1, points: 0.5, grade: true, solution: false -->
```python
# This test should display the max hillslope angle in degrees
print('The maximum hillslope angle in degrees should now be 14.0. My calculated value is {0:.1f}.'.format(max_slope_deg))

```

### Part 4, plot 3: Half the diffusivity

Please create the plot and do the related calculations below.

<!-- grade_id: problem_1_ag_p4_plot4, points: 0, grade: false, solution: true -->
```python
# Reduce the diffusivity by half
hillslope['Half diffusivity'] = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_1_mg_p4_plot4, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: problem_1_ag_p4_plot4_test1, points: 0.5, grade: true, solution: false -->
```python
# This test should display the max hillslope angle in degrees
print('The maximum hillslope angle in degrees should now be 45.0. My calculated value is {0:.1f}.'.format(max_slope_deg))

```

## Part 5: Questions for Problem 1 (2 points)

1. At what value of $x$ (distance from the divide) is the maximum slope?
2. Where is the maximum slope in relation to the crest of the interfluve and the river channel?
3. What does a characteristic timescale mean?
4. What is the value for the characteristic timescale for Part 3? Does it seem reasonable? Why or why not?

<!-- grade_id: problem_1_mg_p5_questions, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

# Problem 2: Incision history of the western Sierra Nevada mountains, California, USA (*7.5 points*)

In this problem we'll continue using our diffusive hillslope equation, but apply it to natural interfluves in mountainous regions. Active mountain ranges typically have poorly developed soils and abundant exposed bedrock.

For this section we will apply our model equation to a real landscape, the western Sierra Nevada mountains in California, USA. We will use a [topographic profile](data/sierras_profile.txt) extracted across an interfluve between two streams draining into the Yuba River and change the landscape uplift rate in the diffusive hillslope equation until we get a reasonable match to the observed profile.

Before we do anything, it is important that you know the location of the topographic profile.

![Topographic profile location](img/Sierras_profile_map.png)
*Figure 1. Shaded relief digital elevation model of the western Sierra Nevada Mountains in California, USA. The line A-A´ is the location of a topographic profile used in Problem 2.*

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable names mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.

**Your score on this problem will be based on following criteria**:

- Reading in the data file
- Calculating hillslope profiles for the geometry of the observed hillslope profile
- Plotting your results
- Creating a goodness-of-fit function for comparing the natural and predicted hillslope profiles
- Finding an uplift rate `U` that best reproduces the observed hillslope profile
- Including comments that explain what most lines in the code do
- Uploading your notebook to your GitHub repository for this week's exercise

## Part 1: Reading the data file (*1 point*)

Now that you know where the profile is located we can read the [data file](data/sierras_profile.txt) using Pandas.
It contains distances from the drainage divide and elevations for the topographic profile A-A´.

To help you in deciding how to read in the data file, the top 5 lines of the data file are below.
```
Distance (m),Observed elevation (m)
-4.8200000e+003,7.4000000e+002
-4.7882718e+003,7.5200000e+002
-4.7565436e+003,7.7100000e+002
-4.7248155e+003,7.9700000e+002
```

- Read the data file into a DataFrame called `sierras` with the given column headings
    - Use 'Distance (m)' as the index column

<!-- grade_id: problem_2_ag_p1_read_data, points: 0, grade: false, solution: true -->
```python
sierras = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_ag_p1_read_data_test1, points: 1, grade: true, solution: false -->
```python
# This test should print the first row of the data file
print("First 5 rows of the data file:\n", sierras.head())

```

## Part 2: Extracting mountain hillslope profile data (*1 point*)

With the data file loaded, we can now calculate an example diffusive hillslope profile using the geometry of the natural hillslope from the Sierra Nevada mountains. Let's start by having a look at the hillslope profile.

![Topographic profile](img/sierras_profile.png)

*Figure 2. Topographic profile across an interfluve between 2 streams draining into the Yuba River, Sierra Nevada mountains, California, USA.*

To make our plot, we need to extract the half-width of the profile and the distances at which the hillslope profile should be calculated.

- Find the hillslope half-width from the data read from the hillslope profile data file
    - This value can be calculated as the absolute value of the minimum of the `sierras` DataFrame index
    - Define this as variable `L_sierras`
- Find the distances at which the profile elevations should be calculated
    - These should be the same distances from the data file
    - Define these as variable `x_sierras`

<!-- grade_id: problem_2_ag_p2_find_vals, points: 0, grade: false, solution: true -->
```python
# Find the half-width of the hillslope profile
L_sierras = None

# Find the distances at which to calculate the hillslope elevations
x_sierras = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_ag_p2_find_vals_test1, points: 1, grade: true, solution: false -->
```python
# This should print the first value of x
print('The first value of x should be -4820.0. My value is {0:.1f}.'.format(x_sierras[0]))

```

## Part 3: A mountain hillslope profile (*1.5 points*)

For this question, you can use your same `hillslope_profile` function from Problem 1.

- Calculate a hillslope profile assuming an uplift rate of 0.5 mm/a, and an appropriate diffusivity for rock of 1.8 m<sup>2</sup>/a
    - The $L$ and $x$ values should be taken from Part 2 of this problem
- Add the calculated hillslope profile as a new column in the `sierras` DataFrame called `'Predicted elevation (m)'`

- Plot the value of `'Predicted elevation (m)'` as a function of the distance along your hillslope profile
    - Use a color of your choice for the line
    - It may be helpful to again increase the figure size to something like 12 inches by 8 inches
- Add axis labels and a plot title
- Again calculate the maximum slope and its value in degrees, the maximum relief, and the characteristic timescale and display those values on the plot
- Add a figure caption in the Markdown cell below the plot explaining what it shows as if it was in a scientific publication

<!-- grade_id: problem_2_ag_p3_plot1, points: 0, grade: false, solution: true -->
```python
# Calculate elevations for natural landscape
sierras['Predicted elevation (m)'] = None

# Calculate the hillslope properties
max_slope = None
max_slope_deg = None
max_relief = None
tau = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_mg_p3_plot1, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: problem_2_ag_p3_plot1_test1, points: 0.5, grade: true, solution: false -->
```python
# This test should display the max hillslope angle in degrees
print('The maximum hillslope angle in degrees should now be 53.2. My calculated value is {0:.1f}.'.format(max_slope_deg))

```

## Part 4: A modified goodness-of-fit function (*1 point*)

Our next task is to create a Python function for calculating a goodness of fit. This will allow us to *quantitatively* compare our diffusive hillslope profiles to the natural hillslope profile from the Sierra Nevada mountains.

For this we will use a slightly different form of the chi-squared function we used last week, since we do not have a standard deviation value for the oberved elevation data. The modified form is

\begin{equation}
  \Large
  \chi^{2} = \sum \frac{(O_{i} - E_{i})^{2}}{E_{i}}
\end{equation}

where $O_{i}$ is the $i$th observed elevation and $E_{i}$ is the $i$th expected elevation. In this case, the $E_{i}$ values come from our hillslope profile function.

- Create a goodness-of-fit function called `chi_squared`

<!-- grade_id: problem_2_ag_p4_function, points: 0, grade: false, solution: true -->
```python
def chi_squared(observed, expected):
    """Returns the chi-squared value for input array data."""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_ag_p4_function_test1, points: 1, grade: true, solution: false -->
```python
# These visible tests with fake data should work
from nose.tools import ok_, assert_equal

# Fake data
obs1 = np.array([1.1, 2.9, 2.6, 3.5, 5.7, 2.8])
exp1 = np.array([1.5, 2.4, 3.6, 1.5, 6.7, 2.6])

obs2 = np.array([1.8, 2.3, 2.2, 3.9, 5.5, 2.4])
exp2 = np.array([1.2, 2.0, 3.9, 1.7, 6.1, 2.9])

# Fake goodness-of-fit values
cs1 = chi_squared(obs1, exp1)
cs2 = chi_squared(obs2, exp2)

# Print results
print("Goodness-of-fit for dataset 1: {0:.4f}.".format(cs1))
print("Goodness-of-fit for dataset 2: {0:.4f}.".format(cs2))

# Check that the chi-squared function works
assert_equal(round(cs1, 4), 3.3199)
assert_equal(round(cs2, 4), 4.0783)
```

## Part 5: Fitting the model to the data (*2 points*)

Your final coding task is to vary some the uplift rate for the hillslope diffusion equation in order to try to match the hillslope profile from the Sierras in Figure 2. You will only be exploring variations in landscape uplift rates (`U`), but the values for the other variables will not change.

- Calculate the diffusive hillslope profile similarly to how it was done in Part 3
    - You can overwrite the values in the `sierras['Predicted elevation (m)'` column with the new calculation
    - **Note**: Because the elevation of the rivers in the topographic profile are at ~750 m above sea level, you will need to shift all the model profile elevations upwards by 750 m to compare them
- Plot your calculated hillslope profile in the same way as in Part 3, with a few minor changes
    - Plot the observed topographic profile from the data file using a black dashed line and the calculated hillslope profile using a color of your choice
    - Adjust the landscape uplift rate `U` until you have a good match between the geometries of the calculated hillslope profile and the observed topographic profile
         - You can use your goodness-of-fit value to guide your selection of `U` values
         - You do not need to find the absolute minimum goodness-of-fit, just the value where the goodness-of-fit does not decrease when you increase or decrease `U` by $\pm 0.02$ mm/a.
    - Add some text to the plot using the `ax.text()` function to display
        1. Your best-fit uplift rate
        2. The goodness-of-fit
    - Note that you do not need to include the other values (slope, relief, etc.) from Part 1 as text on this plot.
- Add a figure caption in the Markdown cell below the plot explaining what it shows as if it was in a scientific publication

<!-- grade_id: problem_2_p5_mg_fit, points: 1.5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_mg_p5_caption, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

## Part 6: Questions for Problem 2 (*1 point*)

1. Discuss the implications of your results for Part 1 for mountain hillslopes in 2-3 sentences. What do the various values you have calculated imply for natural systems? How long might they take to respond to changes in river erosion rates, for example?
2. Does the uplift rate that best fits the observed topographic profile for the Sierra Nevada mountains seem reasonable? Why or why not? What information have you based your answer on?

<!-- grade_id: problem_2_mg_p6_questions, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE